---
title: 'Getting Started with plotly'
description: 'Why plotly? is matplotlib and seaborn not enough?'
---

## Basics 101

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

Welcome to Plotly 101, in this course we will take a look at plotting basic graphs using plotly.

`@instructions`
First let's import few libraries to get started.

`@hint`
import plotly's graph objects as go and also import cufflinks

`@pre_exercise_code`
```{python}
# plotly standard imports graph object and plolty
import plotly
import 

# Cufflinks wrapper on plotly
import 

# Data science imports
import pandas as pd
import numpy as np

# Options for pandas
pd.options.display.max_columns = 30

# Display all cell outputs
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = 'all'

```

`@sample_code`
```{python}
# plotly standard imports graph object and plolty
import plotly
import 

# Cufflinks wrapper on plotly
import 
```

`@solution`
```{python}
# plotly standard imports
import plotly.graph_objs as go
import plotly.plotly as py



```

`@sct`
```{python}

```
